% Running in auto input_syntax mode. Trying TPTP
% Refutation found. Thanks to Tanya!
% SZS status Theorem for horn_example_eq_head_inlined
% SZS output start Proof for horn_example_eq_head_inlined
fof(f1,axiom,(
  ! [X0] : ((p(X0) & q(X0)) => X0 = zero)),
  file('/Users/kondylidou/Developer/Taelja/test/input/horn_example_eq_head_inlined.p',unknown)).
fof(f2,axiom,(
  ! [X0] : p(f(X0))),
  file('/Users/kondylidou/Developer/Taelja/test/input/horn_example_eq_head_inlined.p',unknown)).
fof(f3,axiom,(
  ! [X0] : q(f(X0))),
  file('/Users/kondylidou/Developer/Taelja/test/input/horn_example_eq_head_inlined.p',unknown)).
fof(f4,axiom,(
  ! [X0] : f(X0) = g(X0)),
  file('/Users/kondylidou/Developer/Taelja/test/input/horn_example_eq_head_inlined.p',unknown)).
fof(f5,conjecture,(
  g(a) = zero),
  file('/Users/kondylidou/Developer/Taelja/test/input/horn_example_eq_head_inlined.p',unknown)).
fof(f6,negated_conjecture,(
  ~g(a) = zero),
  inference(negated_conjecture,[status(cth)],[f5])).
fof(f7,plain,(
  zero != g(a)),
  inference(flattening,[],[f6])).
fof(f8,plain,(
  ! [X0] : (X0 = zero | (~p(X0) | ~q(X0)))),
  inference(ennf_transformation,[],[f1])).
fof(f9,plain,(
  ! [X0] : (X0 = zero | ~p(X0) | ~q(X0))),
  inference(flattening,[],[f8])).
fof(f10,plain,(
  ( ! [X0] : (~q(X0) | ~p(X0) | zero = X0) )),
  inference(cnf_transformation,[],[f9])).
fof(f11,plain,(
  ( ! [X0] : (p(f(X0))) )),
  inference(cnf_transformation,[],[f2])).
fof(f12,plain,(
  ( ! [X0] : (q(f(X0))) )),
  inference(cnf_transformation,[],[f3])).
fof(f13,plain,(
  ( ! [X0] : (f(X0) = g(X0)) )),
  inference(cnf_transformation,[],[f4])).
fof(f14,plain,(
  zero != g(a)),
  inference(cnf_transformation,[],[f7])).
fof(f15,plain,(
  ( ! [X0] : (p(g(X0))) )),
  inference(definition_unfolding,[],[f11,f13])).
fof(f16,plain,(
  ( ! [X0] : (q(g(X0))) )),
  inference(definition_unfolding,[],[f12,f13])).
fof(f17,plain,(
  ( ! [X0] : (~p(g(X0)) | zero = g(X0)) )),
  inference(resolution,[],[f10,f16])).
fof(f18,plain,(
  ( ! [X0] : (zero = g(X0)) )),
  inference(forward_subsumption_resolution,[],[f17,f15])).
fof(f19,plain,(
  zero != zero),
  inference(superposition,[],[f14,f18])).
fof(f22,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f19])).
% SZS output end Proof for horn_example_eq_head_inlined
% ------------------------------
% Version: Vampire 5.0.1 (Release build, commit 1b13eaf on 2026-01-18 12:14:50 +0000)
% Linked with Z3 4.14.0.0 3c47fd96cf5645d0c42b2c819d9e9a84380aa721 NOTFOUND
% CaDiCaL version: 2.1.3
% Termination reason: Refutation
% Time elapsed: 0.001 s
% Peak memory usage: 7 MB
% ------------------------------
% ------------------------------
