% Running in auto input_syntax mode. Trying TPTP
% Refutation found. Thanks to Tanya!
% SZS status Theorem for test_havehence_in_eqchain
% SZS output start Proof for test_havehence_in_eqchain
fof(f1,axiom,(
  p(a)),
  file('/Users/kondylidou/Developer/Taelja/test/input/test_havehence_in_eqchain.p',unknown)).
fof(f2,axiom,(
  q(a)),
  file('/Users/kondylidou/Developer/Taelja/test/input/test_havehence_in_eqchain.p',unknown)).
fof(f3,axiom,(
  ! [X0] : (p(X0) => f(X0) = g(X0))),
  file('/Users/kondylidou/Developer/Taelja/test/input/test_havehence_in_eqchain.p',unknown)).
fof(f4,axiom,(
  ! [X0] : (q(X0) => g(X0) = h(X0))),
  file('/Users/kondylidou/Developer/Taelja/test/input/test_havehence_in_eqchain.p',unknown)).
fof(f5,conjecture,(
  s(f(a)) = s(h(a))),
  file('/Users/kondylidou/Developer/Taelja/test/input/test_havehence_in_eqchain.p',unknown)).
fof(f6,negated_conjecture,(
  ~s(f(a)) = s(h(a))),
  inference(negated_conjecture,[status(cth)],[f5])).
fof(f7,plain,(
  s(f(a)) != s(h(a))),
  inference(flattening,[],[f6])).
fof(f8,plain,(
  ! [X0] : (f(X0) = g(X0) | ~p(X0))),
  inference(ennf_transformation,[],[f3])).
fof(f9,plain,(
  ! [X0] : (g(X0) = h(X0) | ~q(X0))),
  inference(ennf_transformation,[],[f4])).
fof(f10,plain,(
  p(a)),
  inference(cnf_transformation,[],[f1])).
fof(f11,plain,(
  q(a)),
  inference(cnf_transformation,[],[f2])).
fof(f12,plain,(
  ( ! [X0] : (~p(X0) | f(X0) = g(X0)) )),
  inference(cnf_transformation,[],[f8])).
fof(f13,plain,(
  ( ! [X0] : (~q(X0) | g(X0) = h(X0)) )),
  inference(cnf_transformation,[],[f9])).
fof(f14,plain,(
  s(f(a)) != s(h(a))),
  inference(cnf_transformation,[],[f7])).
fof(f15,plain,(
  f(a) = g(a)),
  inference(resolution,[],[f12,f10])).
fof(f16,plain,(
  h(a) = g(a)),
  inference(resolution,[],[f13,f11])).
fof(f17,plain,(
  f(a) = h(a)),
  inference(forward_demodulation,[],[f16,f15])).
fof(f18,plain,(
  s(f(a)) != s(f(a))),
  inference(superposition,[],[f14,f17])).
fof(f19,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f18])).
% SZS output end Proof for test_havehence_in_eqchain
% ------------------------------
% Version: Vampire 5.0.1 (Release build, commit 1b13eaf on 2026-01-18 12:14:50 +0000)
% Linked with Z3 4.14.0.0 3c47fd96cf5645d0c42b2c819d9e9a84380aa721 NOTFOUND
% CaDiCaL version: 2.1.3
% Termination reason: Refutation
% Time elapsed: 0.0000 s
% Peak memory usage: 7 MB
% ------------------------------
% ------------------------------
