# Preprocessing class: FSSSSMSSSSSNFFN.
# Configuration: G-E--_302_C18_F1_URBAN_RG_S04BN
# No SInE strategy applied
# Search class: FHUPF-FFSF11-SFFFFFNN
# partial match(1): FHUNF-FFSF11-SFFFFFNN
# Configuration: SAT001_MinMin_p005000_rr_RG
# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_1, plain, (b=a)).
#
#cnf(i_0_5, negated_conjecture, (h(a)!=c)).
#
#cnf(i_0_4, plain, (g(a)!=c)).
#
#cnf(i_0_2, plain, (f(a)=c)).
#
#cnf(i_0_3, plain, (g(X1)=c|f(X1)!=c)).
# Presaturation interreduction done
#
#cnf(i_0_1, plain, (b=a)).
#
#cnf(i_0_2, plain, (f(a)=c)).
#
#cnf(i_0_5, negated_conjecture, (h(a)!=c)).
#
#cnf(i_0_4, plain, (g(a)!=c)).
#
#cnf(i_0_3, plain, (g(X1)=c|f(X1)!=c)).

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(goal, conjecture, h(a)=c, file('/Users/kondylidou/Developer/Taelja/test/input/horn_example_eq_rw_chain.p', goal)).
fof(ax4, axiom, (g(a)=c=>h(a)=c), file('/Users/kondylidou/Developer/Taelja/test/input/horn_example_eq_rw_chain.p', ax4)).
fof(ax3, axiom, ![X1]:((f(X1)=c=>g(X1)=c)), file('/Users/kondylidou/Developer/Taelja/test/input/horn_example_eq_rw_chain.p', ax3)).
fof(ax2, axiom, f(b)=c, file('/Users/kondylidou/Developer/Taelja/test/input/horn_example_eq_rw_chain.p', ax2)).
fof(ax1, axiom, a=b, file('/Users/kondylidou/Developer/Taelja/test/input/horn_example_eq_rw_chain.p', ax1)).
fof(c_0_5, negated_conjecture, h(a)!=c, inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[goal])])).
fof(c_0_6, plain, (g(a)!=c|h(a)=c), inference(fof_nnf,[status(thm)],[inference(fof_nnf,[status(thm)],[ax4])])).
fof(c_0_7, negated_conjecture, h(a)!=c, inference(fof_nnf,[status(thm)],[c_0_5])).
cnf(c_0_8, plain, (h(a)=c|g(a)!=c), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_9, negated_conjecture, (h(a)!=c), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_10, plain, ![X2]:((f(X2)!=c|g(X2)=c)), inference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[ax3])])])).
cnf(c_0_11, plain, (f(b)=c), inference(split_conjunct,[status(thm)],[ax2])).
cnf(c_0_12, plain, (a=b), inference(split_conjunct,[status(thm)],[ax1])).
cnf(c_0_13, plain, (g(a)!=c), inference(sr,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_14, plain, (g(X1)=c|f(X1)!=c), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15, plain, (f(a)=c), inference(rw,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_16, plain, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15])]), ['proof']).
# SZS output end CNFRefutation
