# Preprocessing class: FSSSSMSSSSSNFFN.
# Configuration: G-E--_302_C18_F1_URBAN_RG_S04BN
# No SInE strategy applied
# Search class: FHUNF-FFSF11-SFFFFFNN
# Configuration: SAT001_MinMin_p005000_rr_RG
# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_5, negated_conjecture, (~r(a))).
#
#cnf(i_0_2, plain, (q(b,X1))).
#
#cnf(i_0_4, plain, (r(X1)|~q(f(b),X1))).
#
#cnf(i_0_3, plain, (q(f(X1),X2)|~q(X1,X2))).
# Presaturation interreduction done
#
#cnf(i_0_5, negated_conjecture, (~r(a))).
#
#cnf(i_0_2, plain, (q(b,X1))).
#
#cnf(i_0_4, plain, (r(X1)|~q(f(b),X1))).
#
#cnf(i_0_3, plain, (q(f(X1),X2)|~q(X1,X2))).
#
#cnf(i_0_6, plain, (r(X1))).

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(ax3, axiom, ![X2, X3]:(((q(X2,X3)&p(X3))=>q(f(X2),X3))), file('/Users/kondylidou/Developer/Taelja/test/input/resolution_example_horn_reuse_n1.p', ax3)).
fof(ax1, axiom, ![X1]:(p(X1)), file('/Users/kondylidou/Developer/Taelja/test/input/resolution_example_horn_reuse_n1.p', ax1)).
fof(goal, conjecture, r(a), file('/Users/kondylidou/Developer/Taelja/test/input/resolution_example_horn_reuse_n1.p', goal)).
fof(ax4, axiom, ![X3]:((q(f(b),X3)=>r(X3))), file('/Users/kondylidou/Developer/Taelja/test/input/resolution_example_horn_reuse_n1.p', ax4)).
fof(ax2, axiom, ![X1]:(q(b,X1)), file('/Users/kondylidou/Developer/Taelja/test/input/resolution_example_horn_reuse_n1.p', ax2)).
fof(c_0_5, plain, ![X6, X7]:((~q(X6,X7)|~p(X7)|q(f(X6),X7))), inference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[ax3])])])).
fof(c_0_6, plain, ![X4]:(p(X4)), inference(variable_rename,[status(thm)],[ax1])).
fof(c_0_7, negated_conjecture, ~r(a), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[goal])])).
fof(c_0_8, plain, ![X8]:((~q(f(b),X8)|r(X8))), inference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[ax4])])])).
cnf(c_0_9, plain, (q(f(X1),X2)|~q(X1,X2)|~p(X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_10, plain, (p(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_11, plain, ![X5]:(q(b,X5)), inference(variable_rename,[status(thm)],[ax2])).
fof(c_0_12, negated_conjecture, ~r(a), inference(fof_nnf,[status(thm)],[c_0_7])).
cnf(c_0_13, plain, (r(X1)|~q(f(b),X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14, plain, (q(f(X1),X2)|~q(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_9, c_0_10])])).
cnf(c_0_15, plain, (q(b,X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16, negated_conjecture, (~r(a)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_17, plain, (r(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15])])).
cnf(c_0_18, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_16, c_0_17])]), ['proof']).
# SZS output end CNFRefutation
