# Preprocessing class: FSSSSMSSSSSNFFN.
# Configuration: G-E--_302_C18_F1_URBAN_RG_S04BN
# No SInE strategy applied
# Search class: FHUNF-FFSF00-SFFFFFNN
# Configuration: SAT001_MinMin_p005000_rr_RG
# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_6, negated_conjecture, (~r0(a))).
#
#cnf(i_0_1, plain, (p(a))).
#
#cnf(i_0_5, plain, (r0(X1)|~r2(X1)|~r1(X1))).
#
#cnf(i_0_3, plain, (r1(X1)|~q(X1,b))).
#
#cnf(i_0_4, plain, (r2(X1)|~q(X1,c))).
#
#cnf(i_0_2, plain, (q(X1,X2)|~p(X1))).
# Presaturation interreduction done
#
#cnf(i_0_1, plain, (p(a))).
#
#cnf(i_0_6, negated_conjecture, (~r0(a))).
#
#cnf(i_0_3, plain, (r1(X1)|~q(X1,b))).
#
#cnf(i_0_4, plain, (r2(X1)|~q(X1,c))).
#
#cnf(i_0_5, plain, (r0(X1)|~r2(X1)|~r1(X1))).
#
#cnf(i_0_7, negated_conjecture, (~r2(a)|~r1(a))).
#
#cnf(i_0_8, negated_conjecture, (~r1(a)|~q(a,c))).
#
#cnf(i_0_9, negated_conjecture, (~q(a,c)|~q(a,b))).
#
#cnf(i_0_2, plain, (q(X1,X2)|~p(X1))).
#
#cnf(i_0_10, plain, (q(a,X1))).

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(goal, conjecture, r0(a), file('/Users/kondylidou/Developer/Taelja/test/input/resolution_example_horn_dag.p', goal)).
fof(ax5, axiom, ![X1]:(((r1(X1)&r2(X1))=>r0(X1))), file('/Users/kondylidou/Developer/Taelja/test/input/resolution_example_horn_dag.p', ax5)).
fof(ax4, axiom, ![X1]:((q(X1,c)=>r2(X1))), file('/Users/kondylidou/Developer/Taelja/test/input/resolution_example_horn_dag.p', ax4)).
fof(ax3, axiom, ![X1]:((q(X1,b)=>r1(X1))), file('/Users/kondylidou/Developer/Taelja/test/input/resolution_example_horn_dag.p', ax3)).
fof(ax2, axiom, ![X1, X2]:((p(X1)=>q(X1,X2))), file('/Users/kondylidou/Developer/Taelja/test/input/resolution_example_horn_dag.p', ax2)).
fof(ax1, axiom, p(a), file('/Users/kondylidou/Developer/Taelja/test/input/resolution_example_horn_dag.p', ax1)).
fof(c_0_6, negated_conjecture, ~r0(a), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[goal])])).
fof(c_0_7, negated_conjecture, ~r0(a), inference(fof_nnf,[status(thm)],[c_0_6])).
fof(c_0_8, plain, ![X7]:((~r1(X7)|~r2(X7)|r0(X7))), inference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[ax5])])])).
cnf(c_0_9, negated_conjecture, (~r0(a)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_10, plain, (r0(X1)|~r1(X1)|~r2(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_11, plain, ![X6]:((~q(X6,c)|r2(X6))), inference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[ax4])])])).
cnf(c_0_12, negated_conjecture, (~r2(a)|~r1(a)), inference(spm,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_13, plain, (r2(X1)|~q(X1,c)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_14, plain, ![X5]:((~q(X5,b)|r1(X5))), inference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[ax3])])])).
fof(c_0_15, plain, ![X3, X4]:((~p(X3)|q(X3,X4))), inference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[ax2])])])).
cnf(c_0_16, negated_conjecture, (~r1(a)|~q(a,c)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_17, plain, (r1(X1)|~q(X1,b)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_18, plain, (q(X1,X2)|~p(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_19, plain, (p(a)), inference(split_conjunct,[status(thm)],[ax1])).
cnf(c_0_20, negated_conjecture, (~q(a,c)|~q(a,b)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_21, plain, (q(a,X1)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_22, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_20, c_0_21]), c_0_21])]), ['proof']).
# SZS output end CNFRefutation
