# Preprocessing class: FSSSSMSSSSSNFFN.
# Configuration: G-E--_302_C18_F1_URBAN_RG_S04BN
# No SInE strategy applied
# Search class: FHUNF-FFSF00-SFFFFFNN
# Configuration: SAT001_MinMin_p005000_rr_RG
# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_1, plain, (s(a))).
#
#cnf(i_0_6, negated_conjecture, (~r(a,b))).
#
#cnf(i_0_3, plain, (t(b))).
#
#cnf(i_0_2, plain, (p(X1)|~s(X1))).
#
#cnf(i_0_4, plain, (q(X1)|~t(X1))).
#
#cnf(i_0_5, plain, (r(X1,X2)|~q(X2)|~p(X1))).
# Presaturation interreduction done
#
#cnf(i_0_1, plain, (s(a))).
#
#cnf(i_0_6, negated_conjecture, (~r(a,b))).
#
#cnf(i_0_3, plain, (t(b))).
#
#cnf(i_0_2, plain, (p(X1)|~s(X1))).
#
#cnf(i_0_4, plain, (q(X1)|~t(X1))).
#
#cnf(i_0_5, plain, (r(X1,X2)|~q(X2)|~p(X1))).
#
#cnf(i_0_7, negated_conjecture, (~q(b)|~p(a))).
#
#cnf(i_0_8, negated_conjecture, (~p(a))).

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(goal, conjecture, r(a,b), file('/Users/kondylidou/Developer/Taelja/test/input/resolution_example_horn_general.p', goal)).
fof(ax5, axiom, ![X1, X2]:(((p(X1)&q(X2))=>r(X1,X2))), file('/Users/kondylidou/Developer/Taelja/test/input/resolution_example_horn_general.p', ax5)).
fof(ax4, axiom, ![X2]:((t(X2)=>q(X2))), file('/Users/kondylidou/Developer/Taelja/test/input/resolution_example_horn_general.p', ax4)).
fof(ax3, axiom, t(b), file('/Users/kondylidou/Developer/Taelja/test/input/resolution_example_horn_general.p', ax3)).
fof(ax2, axiom, ![X1]:((s(X1)=>p(X1))), file('/Users/kondylidou/Developer/Taelja/test/input/resolution_example_horn_general.p', ax2)).
fof(ax1, axiom, s(a), file('/Users/kondylidou/Developer/Taelja/test/input/resolution_example_horn_general.p', ax1)).
fof(c_0_6, negated_conjecture, ~r(a,b), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[goal])])).
fof(c_0_7, negated_conjecture, ~r(a,b), inference(fof_nnf,[status(thm)],[c_0_6])).
fof(c_0_8, plain, ![X5, X6]:((~p(X5)|~q(X6)|r(X5,X6))), inference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[ax5])])])).
cnf(c_0_9, negated_conjecture, (~r(a,b)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_10, plain, (r(X1,X2)|~p(X1)|~q(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_11, plain, ![X4]:((~t(X4)|q(X4))), inference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[ax4])])])).
cnf(c_0_12, negated_conjecture, (~q(b)|~p(a)), inference(spm,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_13, plain, (q(X1)|~t(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_14, plain, (t(b)), inference(split_conjunct,[status(thm)],[ax3])).
fof(c_0_15, plain, ![X3]:((~s(X3)|p(X3))), inference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[ax2])])])).
cnf(c_0_16, negated_conjecture, (~p(a)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_14])])).
cnf(c_0_17, plain, (p(X1)|~s(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_18, plain, (s(a)), inference(split_conjunct,[status(thm)],[ax1])).
cnf(c_0_19, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_18])]), ['proof']).
# SZS output end CNFRefutation
