# Preprocessing class: FSSSSMSSSSSNFFN.
# Configuration: G-E--_302_C18_F1_URBAN_RG_S04BN
# No SInE strategy applied
# Search class: FHUPF-FFSF11-SFFFFFNN
# partial match(1): FHUNF-FFSF11-SFFFFFNN
# Configuration: SAT001_MinMin_p005000_rr_RG
# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_4, negated_conjecture, (b!=a)).
#
#cnf(i_0_3, plain, (g(f(b))=g(f(a)))).
#
#cnf(i_0_1, plain, (X1=X2|f(X1)!=f(X2))).
#
#cnf(i_0_2, plain, (X1=X2|g(X1)!=g(X2))).
# Presaturation interreduction done
#
#cnf(i_0_4, negated_conjecture, (b!=a)).
#
#cnf(i_0_3, plain, (g(f(b))=g(f(a)))).
#
#cnf(i_0_1, plain, (X1=X2|f(X1)!=f(X2))).
#
#cnf(i_0_2, plain, (X1=X2|g(X1)!=g(X2))).
#
#cnf(i_0_7, plain, (X1=f(b)|g(X1)!=g(f(a)))).
#
#cnf(i_0_9, plain, (f(b)=f(a))).
#
#cnf(i_0_10, plain, (b=X1|f(a)!=f(X1))).

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(ax2, axiom, ![X1, X2]:((g(X1)=g(X2)=>X1=X2)), file('/Users/kondylidou/Developer/Taelja/test/input/superposition_example_clausal2.p', ax2)).
fof(ax3, axiom, g(f(a))=g(f(b)), file('/Users/kondylidou/Developer/Taelja/test/input/superposition_example_clausal2.p', ax3)).
fof(ax1, axiom, ![X1, X2]:((f(X1)=f(X2)=>X1=X2)), file('/Users/kondylidou/Developer/Taelja/test/input/superposition_example_clausal2.p', ax1)).
fof(goal, conjecture, a=b, file('/Users/kondylidou/Developer/Taelja/test/input/superposition_example_clausal2.p', goal)).
fof(c_0_4, plain, ![X5, X6]:((g(X5)!=g(X6)|X5=X6)), inference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[ax2])])])).
cnf(c_0_5, plain, (X1=X2|g(X1)!=g(X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_6, plain, (g(f(a))=g(f(b))), inference(split_conjunct,[status(thm)],[ax3])).
fof(c_0_7, plain, ![X3, X4]:((f(X3)!=f(X4)|X3=X4)), inference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[ax1])])])).
cnf(c_0_8, plain, (X1=f(b)|g(X1)!=g(f(a))), inference(spm,[status(thm)],[c_0_5, c_0_6])).
fof(c_0_9, negated_conjecture, a!=b, inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[goal])])).
cnf(c_0_10, plain, (X1=X2|f(X1)!=f(X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11, plain, (f(b)=f(a)), inference(er,[status(thm)],[c_0_8])).
fof(c_0_12, negated_conjecture, a!=b, inference(fof_nnf,[status(thm)],[c_0_9])).
cnf(c_0_13, plain, (b=X1|f(a)!=f(X1)), inference(spm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_14, negated_conjecture, (a!=b), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_15, plain, ($false), inference(sr,[status(thm)],[inference(er,[status(thm)],[c_0_13]), c_0_14]), ['proof']).
# SZS output end CNFRefutation
