# Preprocessing class: FSSSSMSSSSSNFFN.
# Configuration: G-E--_302_C18_F1_URBAN_RG_S04BN
# No SInE strategy applied
# Search class: FHUNF-FFSF00-SFFFFFNN
# Configuration: SAT001_MinMin_p005000_rr_RG
# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_4, negated_conjecture, (~r(a))).
#
#cnf(i_0_1, plain, (p(a))).
#
#cnf(i_0_3, plain, (~q(a))).
## Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(f4, conjecture, r(a), file('/Users/kondylidou/Developer/Taelja/test/input/test_nonunit_chain.p', f4)).
fof(f3, axiom, ((p(a)&q(a))=>r(a)), file('/Users/kondylidou/Developer/Taelja/test/input/test_nonunit_chain.p', f3)).
fof(f2, axiom, (p(a)=>q(a)), file('/Users/kondylidou/Developer/Taelja/test/input/test_nonunit_chain.p', f2)).
fof(f1, axiom, p(a), file('/Users/kondylidou/Developer/Taelja/test/input/test_nonunit_chain.p', f1)).
fof(c_0_4, negated_conjecture, ~r(a), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[f4])])).
fof(c_0_5, plain, (~p(a)|~q(a)|r(a)), inference(fof_nnf,[status(thm)],[inference(fof_nnf,[status(thm)],[f3])])).
fof(c_0_6, negated_conjecture, ~r(a), inference(fof_nnf,[status(thm)],[c_0_4])).
fof(c_0_7, plain, (~p(a)|q(a)), inference(fof_nnf,[status(thm)],[inference(fof_nnf,[status(thm)],[f2])])).
cnf(c_0_8, plain, (r(a)|~p(a)|~q(a)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9, plain, (p(a)), inference(split_conjunct,[status(thm)],[f1])).
cnf(c_0_10, negated_conjecture, (~r(a)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_11, plain, (q(a)|~p(a)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12, plain, (~q(a)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_8, c_0_9])]), c_0_10])).
cnf(c_0_13, plain, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_11, c_0_9])]), c_0_12]), ['proof']).
# SZS output end CNFRefutation
