# Preprocessing class: FSSSSMSSSSSNFFN.
# Configuration: G-E--_302_C18_F1_URBAN_RG_S04BN
# No SInE strategy applied
# Search class: FHUNF-FFSF00-SFFFFFNN
# Configuration: SAT001_MinMin_p005000_rr_RG
# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_1, plain, (top)).
#
#cnf(i_0_6, negated_conjecture, (~s(a))).
#
#cnf(i_0_5, plain, (s(X1)|~r2(X1)|~r1(X1))).
#
#cnf(i_0_3, plain, (r1(X1)|~q(X1,b))).
#
#cnf(i_0_2, plain, (q(X1,X2))).
#
#cnf(i_0_3, plain, (r1(X1))).
#
#cnf(i_0_5, plain, (s(X1)|~r2(X1))).
#
#cnf(i_0_4, plain, (r2(X1))).
#
#cnf(i_0_5, plain, (s(X1))).
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(f2, axiom, ![X1, X2]:((top=>q(X1,X2))), file('/Users/kondylidou/Developer/Taelja/test/input/test_instantiations_no_ground.p', f2)).
fof(f3, axiom, ![X1]:((q(X1,b)=>r1(X1))), file('/Users/kondylidou/Developer/Taelja/test/input/test_instantiations_no_ground.p', f3)).
fof(f1, axiom, top, file('/Users/kondylidou/Developer/Taelja/test/input/test_instantiations_no_ground.p', f1)).
fof(f5, axiom, ![X1]:(((r1(X1)&r2(X1))=>s(X1))), file('/Users/kondylidou/Developer/Taelja/test/input/test_instantiations_no_ground.p', f5)).
fof(f4, axiom, ![X1]:((q(X1,X1)=>r2(X1))), file('/Users/kondylidou/Developer/Taelja/test/input/test_instantiations_no_ground.p', f4)).
fof(f6, conjecture, s(a), file('/Users/kondylidou/Developer/Taelja/test/input/test_instantiations_no_ground.p', f6)).
fof(c_0_6, plain, ![X3, X4]:((~top|q(X3,X4))), inference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[f2])])])).
fof(c_0_7, plain, ![X5]:((~q(X5,b)|r1(X5))), inference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[f3])])])).
cnf(c_0_8, plain, (q(X1,X2)|~top), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_9, plain, (top), inference(split_conjunct,[status(thm)],[f1])).
fof(c_0_10, plain, ![X7]:((~r1(X7)|~r2(X7)|s(X7))), inference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[f5])])])).
cnf(c_0_11, plain, (r1(X1)|~q(X1,b)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12, plain, (q(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_8, c_0_9])])).
fof(c_0_13, plain, ![X6]:((~q(X6,X6)|r2(X6))), inference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[f4])])])).
fof(c_0_14, negated_conjecture, ~s(a), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[f6])])).
cnf(c_0_15, plain, (s(X1)|~r1(X1)|~r2(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16, plain, (r1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_11, c_0_12])])).
cnf(c_0_17, plain, (r2(X1)|~q(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_18, negated_conjecture, ~s(a), inference(fof_nnf,[status(thm)],[c_0_14])).
cnf(c_0_19, plain, (s(X1)|~r2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_15, c_0_16])])).
cnf(c_0_20, plain, (r2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_17, c_0_12])])).
cnf(c_0_21, negated_conjecture, (~s(a)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_22, plain, (s(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_19, c_0_20])])).
cnf(c_0_23, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_21, c_0_22])]), ['proof']).
# SZS output end CNFRefutation
