# Preprocessing class: FSSSSMSSSSSNFFN.
# Configuration: G-E--_302_C18_F1_URBAN_RG_S04BN
# No SInE strategy applied
# Search class: FHUSF-FFSF11-SFFFFFNN
# Configuration: SAT001_MinMin_p005000_rr_RG
# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_5, negated_conjecture, (~r(a))).
#
#cnf(i_0_1, plain, (h(a)=k(a))).
#
#cnf(i_0_4, plain, (r(X1)|~q(X1)|~p(X1))).
#
#cnf(i_0_2, plain, (p(X1)|h(X1)!=k(X1))).
#
#cnf(i_0_3, plain, (q(X1)|h(X1)!=k(X1))).
# Presaturation interreduction done
#
#cnf(i_0_1, plain, (h(a)=k(a))).
#
#cnf(i_0_2, plain, (p(X1)|h(X1)!=k(X1))).
#
#cnf(i_0_5, negated_conjecture, (~r(a))).
#
#cnf(i_0_3, plain, (q(X1)|h(X1)!=k(X1))).
#
#cnf(i_0_4, plain, (r(X1)|~q(X1)|~p(X1))).
#
#cnf(i_0_6, negated_conjecture, (~q(a)|~p(a))).
#
#cnf(i_0_7, negated_conjecture, (~p(a))).

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(f5, conjecture, r(a), file('/Users/kondylidou/Developer/Taelja/test/input/test_prelemmatize_sym_eq.p', f5)).
fof(f4, axiom, ![X1]:(((p(X1)&q(X1))=>r(X1))), file('/Users/kondylidou/Developer/Taelja/test/input/test_prelemmatize_sym_eq.p', f4)).
fof(f3, axiom, ![X1]:((k(X1)=h(X1)=>q(X1))), file('/Users/kondylidou/Developer/Taelja/test/input/test_prelemmatize_sym_eq.p', f3)).
fof(f1, axiom, h(a)=k(a), file('/Users/kondylidou/Developer/Taelja/test/input/test_prelemmatize_sym_eq.p', f1)).
fof(f2, axiom, ![X1]:((h(X1)=k(X1)=>p(X1))), file('/Users/kondylidou/Developer/Taelja/test/input/test_prelemmatize_sym_eq.p', f2)).
fof(c_0_5, negated_conjecture, ~r(a), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[f5])])).
fof(c_0_6, negated_conjecture, ~r(a), inference(fof_nnf,[status(thm)],[c_0_5])).
fof(c_0_7, plain, ![X4]:((~p(X4)|~q(X4)|r(X4))), inference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[f4])])])).
cnf(c_0_8, negated_conjecture, (~r(a)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_9, plain, (r(X1)|~p(X1)|~q(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_10, plain, ![X3]:((k(X3)!=h(X3)|q(X3))), inference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[f3])])])).
cnf(c_0_11, negated_conjecture, (~q(a)|~p(a)), inference(spm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_12, plain, (q(X1)|k(X1)!=h(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_13, plain, (h(a)=k(a)), inference(split_conjunct,[status(thm)],[f1])).
fof(c_0_14, plain, ![X2]:((h(X2)!=k(X2)|p(X2))), inference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[f2])])])).
cnf(c_0_15, negated_conjecture, (~p(a)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_13])])).
cnf(c_0_16, plain, (p(X1)|h(X1)!=k(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_17, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_13])]), ['proof']).
# SZS output end CNFRefutation
