# Preprocessing class: FSSSSMSSSSSNFFN.
# Configuration: G-E--_302_C18_F1_URBAN_RG_S04BN
# No SInE strategy applied
# Search class: FHUSF-FFSF11-SFFFFFNN
# Configuration: SAT001_MinMin_p005000_rr_RG
# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_1, plain, (p(a))).
#
#cnf(i_0_5, negated_conjecture, (s(h(a))!=s(f(a)))).
#
#cnf(i_0_2, plain, (q(a))).
#
#cnf(i_0_3, plain, (f(X1)=g(X1)|~p(X1))).
#
#cnf(i_0_4, plain, (h(X1)=g(X1)|~q(X1))).
# Presaturation interreduction done
#
#cnf(i_0_1, plain, (p(a))).
#
#cnf(i_0_2, plain, (q(a))).
#
#cnf(i_0_5, negated_conjecture, (s(h(a))!=s(f(a)))).
#
#cnf(i_0_3, plain, (f(X1)=g(X1)|~p(X1))).
#
#cnf(i_0_4, plain, (h(X1)=g(X1)|~q(X1))).
#
#cnf(i_0_6, negated_conjecture, (s(f(a))!=s(g(a)))).

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(f5, conjecture, s(f(a))=s(h(a)), file('/Users/kondylidou/Developer/Taelja/test/input/test_havehence_in_eqchain.p', f5)).
fof(f4, axiom, ![X1]:((q(X1)=>g(X1)=h(X1))), file('/Users/kondylidou/Developer/Taelja/test/input/test_havehence_in_eqchain.p', f4)).
fof(f2, axiom, q(a), file('/Users/kondylidou/Developer/Taelja/test/input/test_havehence_in_eqchain.p', f2)).
fof(f3, axiom, ![X1]:((p(X1)=>f(X1)=g(X1))), file('/Users/kondylidou/Developer/Taelja/test/input/test_havehence_in_eqchain.p', f3)).
fof(f1, axiom, p(a), file('/Users/kondylidou/Developer/Taelja/test/input/test_havehence_in_eqchain.p', f1)).
fof(c_0_5, negated_conjecture, s(f(a))!=s(h(a)), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[f5])])).
fof(c_0_6, negated_conjecture, s(f(a))!=s(h(a)), inference(fof_nnf,[status(thm)],[c_0_5])).
fof(c_0_7, plain, ![X3]:((~q(X3)|g(X3)=h(X3))), inference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[f4])])])).
cnf(c_0_8, negated_conjecture, (s(f(a))!=s(h(a))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_9, plain, (g(X1)=h(X1)|~q(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_10, plain, (q(a)), inference(split_conjunct,[status(thm)],[f2])).
fof(c_0_11, plain, ![X2]:((~p(X2)|f(X2)=g(X2))), inference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[f3])])])).
cnf(c_0_12, negated_conjecture, (s(f(a))!=s(g(a))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_9]), c_0_10])])).
cnf(c_0_13, plain, (f(X1)=g(X1)|~p(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_14, plain, (p(a)), inference(split_conjunct,[status(thm)],[f1])).
cnf(c_0_15, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_14])]), ['proof']).
# SZS output end CNFRefutation
