% Running in auto input_syntax mode. Trying TPTP
% Refutation found. Thanks to Tanya!
% SZS status Theorem for resolution_example_horn_reuse_inlined
% SZS output start Proof for resolution_example_horn_reuse_inlined
fof(f1,axiom,(
  ! [X0] : p(X0)),
  file('/Users/kondylidou/Developer/Taelja/test/input/resolution_example_horn_reuse_inlined.p',unknown)).
fof(f2,axiom,(
  ! [X0] : q(b,X0)),
  file('/Users/kondylidou/Developer/Taelja/test/input/resolution_example_horn_reuse_inlined.p',unknown)).
fof(f3,axiom,(
  ! [X0,X1] : ((q(X0,X1) & p(X1)) => q(f(X0),X1))),
  file('/Users/kondylidou/Developer/Taelja/test/input/resolution_example_horn_reuse_inlined.p',unknown)).
fof(f4,conjecture,(
  q(f(f(b)),a)),
  file('/Users/kondylidou/Developer/Taelja/test/input/resolution_example_horn_reuse_inlined.p',unknown)).
fof(f5,negated_conjecture,(
  ~q(f(f(b)),a)),
  inference(negated_conjecture,[status(cth)],[f4])).
fof(f6,plain,(
  ~q(f(f(b)),a)),
  inference(flattening,[],[f5])).
fof(f7,plain,(
  ! [X0,X1] : (q(f(X0),X1) | (~q(X0,X1) | ~p(X1)))),
  inference(ennf_transformation,[],[f3])).
fof(f8,plain,(
  ! [X0,X1] : (q(f(X0),X1) | ~q(X0,X1) | ~p(X1))),
  inference(flattening,[],[f7])).
fof(f9,plain,(
  ( ! [X0] : (p(X0)) )),
  inference(cnf_transformation,[],[f1])).
fof(f10,plain,(
  ( ! [X0] : (q(b,X0)) )),
  inference(cnf_transformation,[],[f2])).
fof(f11,plain,(
  ( ! [X0,X1] : (q(f(X0),X1) | ~q(X0,X1) | ~p(X1)) )),
  inference(cnf_transformation,[],[f8])).
fof(f12,plain,(
  ~q(f(f(b)),a)),
  inference(cnf_transformation,[],[f6])).
fof(f13,plain,(
  ~q(f(b),a) | ~p(a)),
  inference(resolution,[],[f11,f12])).
fof(f14,plain,(
  ~q(f(b),a)),
  inference(forward_subsumption_resolution,[],[f13,f9])).
fof(f15,plain,(
  ~q(b,a) | ~p(a)),
  inference(resolution,[],[f14,f11])).
fof(f16,plain,(
  ~p(a)),
  inference(forward_subsumption_resolution,[],[f15,f10])).
fof(f17,plain,(
  $false),
  inference(forward_subsumption_resolution,[],[f16,f9])).
% SZS output end Proof for resolution_example_horn_reuse_inlined
% ------------------------------
% Version: Vampire 5.0.1 (Release build, commit 1b13eaf on 2026-01-18 12:14:50 +0000)
% Linked with Z3 4.14.0.0 3c47fd96cf5645d0c42b2c819d9e9a84380aa721 NOTFOUND
% CaDiCaL version: 2.1.3
% Termination reason: Refutation
% Time elapsed: 0.0000 s
% Peak memory usage: 7 MB
% ------------------------------
% ------------------------------
