% Running in auto input_syntax mode. Trying TPTP
% Refutation found. Thanks to Tanya!
% SZS status Theorem for superposition_example_nonground_lemma
% SZS output start Proof for superposition_example_nonground_lemma
fof(f1,axiom,(
  ! [X0] : f(X0) = X0),
  file('/Users/kondylidou/Developer/Taelja/test/input/superposition_example_nonground_lemma.p',unknown)).
fof(f2,axiom,(
  ! [X0] : g(X0) = f(X0)),
  file('/Users/kondylidou/Developer/Taelja/test/input/superposition_example_nonground_lemma.p',unknown)).
fof(f3,axiom,(
  ! [X0] : (g(X0) = X0 => p(g(X0)))),
  file('/Users/kondylidou/Developer/Taelja/test/input/superposition_example_nonground_lemma.p',unknown)).
fof(f4,conjecture,(
  p(g(a))),
  file('/Users/kondylidou/Developer/Taelja/test/input/superposition_example_nonground_lemma.p',unknown)).
fof(f5,negated_conjecture,(
  ~p(g(a))),
  inference(negated_conjecture,[status(cth)],[f4])).
fof(f6,plain,(
  ~p(g(a))),
  inference(flattening,[],[f5])).
fof(f7,plain,(
  ! [X0] : (p(g(X0)) | g(X0) != X0)),
  inference(ennf_transformation,[],[f3])).
fof(f8,plain,(
  ( ! [X0] : (f(X0) = X0) )),
  inference(cnf_transformation,[],[f1])).
fof(f9,plain,(
  ( ! [X0] : (f(X0) = g(X0)) )),
  inference(cnf_transformation,[],[f2])).
fof(f10,plain,(
  ( ! [X0] : (g(X0) != X0 | p(g(X0))) )),
  inference(cnf_transformation,[],[f7])).
fof(f11,plain,(
  ~p(g(a))),
  inference(cnf_transformation,[],[f6])).
fof(f12,plain,(
  ( ! [X0] : (g(X0) = X0) )),
  inference(definition_unfolding,[],[f8,f9])).
fof(f13,plain,(
  ~p(a)),
  inference(superposition,[],[f11,f12])).
fof(f14,plain,(
  ( ! [X0] : (X0 != X0 | p(X0)) )),
  inference(superposition,[],[f10,f12])).
fof(f15,plain,(
  ( ! [X0] : (p(X0)) )),
  inference(trivial_inequality_removal,[],[f14])).
fof(f17,plain,(
  $false),
  inference(resolution,[],[f15,f13])).
% SZS output end Proof for superposition_example_nonground_lemma
% ------------------------------
% Version: Vampire 5.0.1 (Release build, commit 1b13eaf on 2026-01-18 12:14:50 +0000)
% Linked with Z3 4.14.0.0 3c47fd96cf5645d0c42b2c819d9e9a84380aa721 NOTFOUND
% CaDiCaL version: 2.1.3
% Termination reason: Refutation
% Time elapsed: 0.0000 s
% Peak memory usage: 7 MB
% ------------------------------
% ------------------------------
