% Running in auto input_syntax mode. Trying TPTP
% Refutation found. Thanks to Tanya!
% SZS status Theorem for test_prelemmatize_sym_eq
% SZS output start Proof for test_prelemmatize_sym_eq
fof(f1,axiom,(
  h(a) = k(a)),
  file('/Users/kondylidou/Developer/Taelja/test/input/test_prelemmatize_sym_eq.p',unknown)).
fof(f2,axiom,(
  ! [X0] : (h(X0) = k(X0) => p(X0))),
  file('/Users/kondylidou/Developer/Taelja/test/input/test_prelemmatize_sym_eq.p',unknown)).
fof(f3,axiom,(
  ! [X0] : (k(X0) = h(X0) => q(X0))),
  file('/Users/kondylidou/Developer/Taelja/test/input/test_prelemmatize_sym_eq.p',unknown)).
fof(f4,axiom,(
  ! [X0] : ((p(X0) & q(X0)) => r(X0))),
  file('/Users/kondylidou/Developer/Taelja/test/input/test_prelemmatize_sym_eq.p',unknown)).
fof(f5,conjecture,(
  r(a)),
  file('/Users/kondylidou/Developer/Taelja/test/input/test_prelemmatize_sym_eq.p',unknown)).
fof(f6,negated_conjecture,(
  ~r(a)),
  inference(negated_conjecture,[status(cth)],[f5])).
fof(f7,plain,(
  ~r(a)),
  inference(flattening,[],[f6])).
fof(f8,plain,(
  ! [X0] : (p(X0) | h(X0) != k(X0))),
  inference(ennf_transformation,[],[f2])).
fof(f9,plain,(
  ! [X0] : (q(X0) | h(X0) != k(X0))),
  inference(ennf_transformation,[],[f3])).
fof(f10,plain,(
  ! [X0] : (r(X0) | (~p(X0) | ~q(X0)))),
  inference(ennf_transformation,[],[f4])).
fof(f11,plain,(
  ! [X0] : (r(X0) | ~p(X0) | ~q(X0))),
  inference(flattening,[],[f10])).
fof(f12,plain,(
  h(a) = k(a)),
  inference(cnf_transformation,[],[f1])).
fof(f13,plain,(
  ( ! [X0] : (h(X0) != k(X0) | p(X0)) )),
  inference(cnf_transformation,[],[f8])).
fof(f14,plain,(
  ( ! [X0] : (h(X0) != k(X0) | q(X0)) )),
  inference(cnf_transformation,[],[f9])).
fof(f15,plain,(
  ( ! [X0] : (~q(X0) | ~p(X0) | r(X0)) )),
  inference(cnf_transformation,[],[f11])).
fof(f16,plain,(
  ~r(a)),
  inference(cnf_transformation,[],[f7])).
fof(f17,plain,(
  h(a) != h(a) | p(a)),
  inference(superposition,[],[f13,f12])).
fof(f18,plain,(
  p(a)),
  inference(trivial_inequality_removal,[],[f17])).
fof(f19,plain,(
  h(a) != h(a) | q(a)),
  inference(superposition,[],[f14,f12])).
fof(f20,plain,(
  q(a)),
  inference(trivial_inequality_removal,[],[f19])).
fof(f21,plain,(
  ~p(a) | r(a)),
  inference(resolution,[],[f20,f15])).
fof(f22,plain,(
  r(a)),
  inference(forward_subsumption_resolution,[],[f21,f18])).
fof(f23,plain,(
  $false),
  inference(forward_subsumption_resolution,[],[f22,f16])).
% SZS output end Proof for test_prelemmatize_sym_eq
% ------------------------------
% Version: Vampire 5.0.1 (Release build, commit 1b13eaf on 2026-01-18 12:14:50 +0000)
% Linked with Z3 4.14.0.0 3c47fd96cf5645d0c42b2c819d9e9a84380aa721 NOTFOUND
% CaDiCaL version: 2.1.3
% Termination reason: Refutation
% Time elapsed: 0.0000 s
% Peak memory usage: 7 MB
% ------------------------------
% ------------------------------
